Прежде чем проверять задания:
- **Перезапустите ядро** (**restart the kernel**) (В меню, выбрать Ядро (Kernel) $\rightarrow$ Перезапустить (Restart)
- Затем **Выполнить все ячейки**  **run all cells** (В меню, выбрать Ячейка (Cell) $\rightarrow$ Запустить все (Run All).

Убедитесь, что заполнены все ячейки с комментарием "НАЧАЛО ВАШЕГО РЕШЕНИЯ".

После ячеек с заданием следуют ячейки с проверкой с помощью assert.

Если в коде есть ошибки, assert выведет уведомление об ошибке.

Если в коде нет ошибок, assert отработает без вывода дополнительной информации.

---

## Импорт библиотек

```python
import numpy as np
import sympy as sp
import scipy.spatial.distance as ds
```

# Задания на векторы

## Задание 1 (1 балл)

Напишите код функции для вычисления L2-нормы одномерный np.array вектора с помощью numpy:

<!-- grade_id: l2_norm, points: 0, grade: false, solution: true -->
```python
def l2_norm(vector):

    # НАЧАЛО ВАШЕГО РЕШЕНИЯ
    length = np.linalg.norm(vector)
    return length
    # КОНЕЦ ВАШЕГО РЕШЕНИЯ
   
```

<!-- grade_id: l2_norm_test, points: 1, grade: true, solution: false -->
```python
# Тест 1: Вектор (1, 0, 0) имеет L2-норму 1.
vector1 = np.array([1, 0, 0])
assert l2_norm(vector1) == 1.0

# Тест 2: Вектор (0, 1, 0) имеет L2-норму 1.
vector2 = np.array([0, 1, 0])
assert l2_norm(vector2) == 1.0

# Тест 3: Вектор (1, 2, 3, 4, 5, 6, 7, 8, 9) имеет L2-норму sqrt(285).
vector3 = np.array([1, 2, 3, 4, 5, 6, 7, 8, 9])
assert l2_norm(vector3) == np.sqrt(285)

# Тест 3: Вектор (0, 0, 1) имеет L2-норму 1.
vector3 = np.array([0, 0, 1])
assert l2_norm(vector3) == 1.0

# Тест 4: Вектор (3, 4) имеет L2-норму 5.
vector4 = np.array([3, 4])
assert l2_norm(vector4) == 5.0

# Тест 5: Вектор (-3, -4) имеет L2-норму 5.
vector5 = np.array([-3, -4])
assert l2_norm(vector5) == 5.0

# Тест 6: Вектор (1, 1, 1) имеет L2-норму sqrt(3).
vector6 = np.array([1, 1, 1])
assert l2_norm(vector6) == np.sqrt(3)

# Тест 7: Вектор (0, 0, 0) имеет L2-норму 0.
vector7 = np.array([0, 0, 0])
assert l2_norm(vector7) == 0.0

# Тест 8: Вектор (1, 2, 3, 4, 5) имеет L2-норму sqrt(55).
vector8 = np.array([1, 2, 3, 4, 5])
assert l2_norm(vector8) == np.sqrt(55)

# Тест 9: Вектор (1, 2, 3, 4, 5, 6) имеет L2-норму sqrt(91).
vector9 = np.array([1, 2, 3, 4, 5, 6])
assert l2_norm(vector9) == np.sqrt(91)

# Тест 10: Вектор (1, 2, 3, 4, 5, 6, 7, 8, 9) имеет L2-норму sqrt(285).
vector10 = np.array([1, 2, 3, 4, 5, 6, 7, 8, 9])
assert l2_norm(vector10) == np.sqrt(285)
```

---
## Задание 2 (1 балл)

Напишите код функции, которая на вход принимает два одномерных np.array массива чисел и производит их сложение с помощью numpy. Если длины векторов не равны, функция должна вернуть None

<!-- grade_id: add_vectors, points: 0, grade: false, solution: true -->
```python
import numpy as np
def add_vectors(arr1, arr2):
    if len(arr1) != len(arr2):
        return None
    else:
        return arr1 + arr2
```

<!-- grade_id: add_vectors_test, points: 1, grade: true, solution: false -->
```python
arr1 = np.array([1, 2, 3])
arr2 = np.array([4, 5, 6])
assert np.array_equal(add_vectors(arr1, arr2), np.array([5, 7, 9]))

arr1 = np.array([1.5, 2.0, 3.5])
arr2 = np.array([4.2, 5.7, 6.1])
assert np.allclose(add_vectors(arr1, arr2), np.array([5.7, 7.7, 9.6]))

arr1 = np.array([])
arr2 = np.array([])
assert np.array_equal(add_vectors(arr1, arr2), np.array([]))

arr1 = np.array([-1, 2, -3])
arr2 = np.array([4, -5, 6])
assert np.array_equal(add_vectors(arr1, arr2), np.array([3, -3, 3]))

arr1 = np.array([1, 2, 3])
arr2 = np.array([4])
assert np.array_equal(add_vectors(arr1, arr2), None)

arr1 = np.array([1, 2, 3])
arr2 = np.array([4, 5])
assert np.array_equal(add_vectors(arr1, arr2), None)

arr1 = np.array([1, 2, 3, 4, 5])
arr2 = np.array([2, 4, 6, 8, 10])
assert np.array_equal(add_vectors(arr1, arr2), np.array([3, 6, 9, 12, 15]))

arr1 = np.array([1, 2, 3])
arr2 = np.array([0, 0, 0])
assert np.array_equal(add_vectors(arr1, arr2), arr1)

arr1 = np.array([1, 2, 3])
arr2 = np.array([-1, -2, -3])
assert np.array_equal(add_vectors(arr1, arr2), np.array([0, 0, 0]))

arr1 = np.array([1, 2, 3])
arr2 = np.array([1, 2, 3])
assert np.array_equal(add_vectors(arr1, arr2), np.array([2, 4, 6]))
```

---
## Задание 3 (1 балл)

Напишите код функции, которая на вход принимает два одномерных np.array массива чисел и вычисляет их скалярное произведение с помощью numpy. Если длины векторов не равны или векторы пустые, необходимо вернуть None

<!-- grade_id: scalar_product, points: 0, grade: false, solution: true -->
```python
def scalar_product(arr1, arr2):

    # НАЧАЛО ВАШЕГО РЕШЕНИЯ
    if len(arr1) != len(arr2):
        return None
    elif len(arr1) == len(arr2) == 0:
        return None
    else:
        dot_product = np.dot(arr1, arr2)
        return dot_product

    # КОНЕЦ ВАШЕГО РЕШЕНИЯ
    

```

<!-- grade_id: scalar_product_test, points: 1, grade: true, solution: false -->
```python
a = np.array([1, 2, 3])
b = np.array([4, 5, 6])
assert scalar_product(a, b) == 32

a = np.array([1, 0, 0])
b = np.array([0, 1, 0])
assert scalar_product(a, b) == 0

a = np.array([1, 2, 3, 4, 5])
b = np.array([5, 4, 3, 2, 1])
assert scalar_product(a, b) == 35

a = np.array([1, 2])
b = np.array([4, 5, 6])
assert scalar_product(a, b) is None

a = np.array([-1, 2, -3])
b = np.array([4, -5, 6])
assert scalar_product(a, b) == -32

a = np.array([0, 0, 0])
b = np.array([4, 5, 6])
assert scalar_product(a, b) == 0

a = np.array([1, 2, 3])
b = np.array([4, 5])
assert scalar_product(a, b) is None

a = np.array([])
b = np.array([4, 5, 6])
assert scalar_product(a, b) is None

a = np.array([])
b = np.array([])
assert scalar_product(a, b) is None

a = np.array([1])
b = np.array([4])
assert scalar_product(a, b) == 4
```

---
## Задание 4 (1 балл)

Напишите код функции, которая на вход принимает два одномерных np.array массива чисел и вычисляет косинусное расстояние между ними с помощью numpy. Если длины векторов не равны, или один из векторов пустой, необходимо вернуть None. В случае, если один из векторов нулевой, также нужно вернуть None

<!-- grade_id: cosine_distance, points: 0, grade: false, solution: true -->
```python
def cosine_distance(arr1, arr2):

    # НАЧАЛО ВАШЕГО РЕШЕНИЯ
    
    # Проверяем равны ли вектора, если нет возвращаем None
    if len(arr1) != len(arr2):
        return None
    # Проверяем есть ли пустые вектора, если есть возвращаем None
    elif len(arr1) == len(arr2) == 0:
        return None
    # Проверяем на нулевой вектор
    elif len(set(arr1)) == 1:
        if arr1[0] == 0:
            return None
        else:
            cosin = 1 - ((np.dot(arr1, arr2))/((np.linalg.norm(arr1)) * (np.linalg.norm(arr2))))
            return cosin
        
    elif len(set(arr2)) == 1:
        if arr2[0] == 0:
            return None
        else:
            cosin = 1 - ((np.dot(arr1, arr2))/((np.linalg.norm(arr1)) * (np.linalg.norm(arr2))))
            return cosin
    # Высчитываем косинусное расстояние       
    else:
        cosin = 1 - ((np.dot(arr1, arr2))/((np.linalg.norm(arr1)) * (np.linalg.norm(arr2))))
        return cosin
    # КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

<!-- grade_id: cosine_distance_test, points: 1, grade: true, solution: false -->
```python
assert np.isclose(cosine_distance(np.array([1, 2, 3]), np.array([4, 5, 6])), 0.025368153802923787)
assert cosine_distance(np.array([1, 2, 3]), np.array([4, 5])) is None
assert cosine_distance(np.array([1, 2, 3]), np.array([])) is None
assert np.isclose(cosine_distance(np.array([0, 1, 0, 1]), np.array([1, 0, 1, 0])), 1.0)
assert np.isclose(cosine_distance(np.array([1, 0, 1]), np.array([-1, 0, -1])), 2.0)
assert np.isclose(cosine_distance(np.array([3, 4]), np.array([6, 8])), 0.0)
assert cosine_distance(np.array([]), np.array([1, 2, 3])) is None
assert cosine_distance(np.array([]), np.array([])) is None
assert cosine_distance(np.array([1, 2, 3]), np.array([0, 0, 0])) is None
assert cosine_distance(np.array([0, 0, 0]), np.array([1, 2, 3])) is None
```

# Задания на матрицы

---
## Задание 5 (1 балл)

Напишите функцию, которая на вход принимает два np.array массива чисел и производит их поэлементное сложение с помощью numpy. Если размеры матриц не совпадают или массивы пустые (не содержат элементов), вернуть None

<!-- grade_id: add_arrays, points: 0, grade: false, solution: true -->
```python
def is_empty(arr):

    return arr.size == 0



def add_arrays(arr1, arr2):

    # НАЧАЛО ВАШЕГО РЕШЕНИЯ

    if arr1.size != arr2.size:
        return None
    elif arr1.size == 0 or arr2.size == 0:
        return None
    else:
        arr_sum = arr1 + arr2
        return arr_sum



    # КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

<!-- grade_id: add_arrays_test, points: 1, grade: true, solution: false -->
```python
arr1 = np.array([1, 2, 3])
arr2 = np.array([4, 5, 6])
assert np.array_equal(add_arrays(arr1, arr2), np.array([5, 7, 9]))

arr1 = np.array([1, 2, 3])
arr2 = np.array([])
assert add_arrays(arr1, arr2) is None

arr1 = np.array([1.1, 2.2, 3.3])
arr2 = np.array([4.4, 5.5, 6.6])
assert np.allclose(add_arrays(arr1, arr2), np.array([5.5, 7.7, 9.9]))

arr1 = np.array([1, 2, 3])
arr2 = np.array([4, 5, 6])
assert np.array_equal(add_arrays(arr1, arr2), add_arrays(arr2, arr1))

arr1 = np.array([1, 2, 3])
arr2 = np.array([4, 5, 6])
arr3 = np.array([7, 8, 9])
assert np.array_equal(add_arrays(add_arrays(arr1, arr2), arr3), add_arrays(arr1, add_arrays(arr2, arr3)))

arr1 = np.array([1, 2, 3])
arr2 = np.array([4, 5, 6])
assert np.array_equal(add_arrays(arr1, arr2), add_arrays(arr1, arr2))

arr1 = np.array([])
arr2 = np.array([1, 2, 3])
assert add_arrays(arr1, arr2) is None

arr1 = np.array([])
arr2 = np.array([[], []])
assert add_arrays(arr1, arr2) is None

arr1 = np.random.rand(1000)
arr2 = np.random.rand(1000)
assert np.allclose(add_arrays(arr1, arr2), arr1 + arr2)
```

---
## Задание 6 (1 балл)

Напишите функцию, которая на вход принимает два np.array массива чисел и производит их умножение с помощью numpy. Должны соблюдаться правила произведения матриц. Если правила не соблюдаются, функция должна вернуть None

<!-- grade_id: matrix_multiply, points: 0, grade: false, solution: true -->
```python
import numpy as np

def matrix_multiply(a, b):

    # НАЧАЛО ВАШЕГО РЕШЕНИЯ
    
    # Никак не мог понять как сделать в а4 и b4, выдавало ошибку
    # После пришла идея сделать из a4 = np.array([1, 2, 3]) a4 = np.array([[1, 2, 3]])
    # Т.е. доавил квадратные скобки для удобства
    a = a * [[1]]
    b = b * [[1]]
    
    if len(a[0]) == b.shape[0]:
        dot_product = np.dot(a, b)
        return dot_product
    else:
        return None

    # КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

```python
a2 = np.array([
    [1, 2, 3], 
    [4, 5, 6]
    ])
b2 = np.array([
    [7, 8], 
    [9, 10], 
    [11, 12]
    ])


a2.shape[0]
```

Output:
```
2
```

<!-- grade_id: matrix_multiply_test, points: 1, grade: true, solution: false -->
```python
a1 = np.array([[1, 2], [3, 4]])
b1 = np.array([[5, 6], [7, 8]])
c1 = np.array([[19, 22], [43, 50]])
assert np.array_equal(matrix_multiply(a1, b1), c1)

a2 = np.array([[1, 2, 3], [4, 5, 6]])
b2 = np.array([[7, 8], [9, 10], [11, 12]])
c2 = np.array([[58, 64], [139, 154]])
assert np.array_equal(matrix_multiply(a2, b2), c2)

a3 = np.array([[1, 2, 3], [4, 5, 6]])
b3 = np.array([[7, 8], [9, 10]])
assert matrix_multiply(a3, b3) is None

a4 = np.array([1, 2, 3])
b4 = np.array([4, 5, 6])
assert matrix_multiply(a4, b4) is None

a5 = np.array([[1, 2], [3, 4]])
b5 = np.array([[5], [6]])
c5 = np.array([[17], [39]])
assert np.array_equal(matrix_multiply(a5, b5), c5)

a6 = np.array([[1, 2, 3]])
b6 = np.array([[4], [5], [6]])
c6 = np.array([[32]])
assert np.array_equal(matrix_multiply(a6, b6), c6)

a7 = np.array([[]])
b7 = np.array([[1]])
assert matrix_multiply(a7, b7) is None

a8 = np.array([[]])
b8 = np.array([[]])
assert matrix_multiply(a8, b8) is None
```

---
## Задание 7 (1 балл)

Напишите функцию, которая на вход принимает np.array массив чисел и вычисляет определитель с помощью numpy. Если матрица не квадратная, функция должна вернуть None

<!-- grade_id: determinant, points: 0, grade: false, solution: true -->
```python
def determinant(arr):
 
    # НАЧАЛО ВАШЕГО РЕШЕНИЯ
    shape_arr = arr.shape
    transpose_arr = arr.transpose()
    shape_arr2 = transpose_arr.shape
    if shape_arr != shape_arr2:
        return None
    elif len(arr) == 0:
        return None 
    else:
        det_arr = np.linalg.det(arr)
        return det_arr
    # КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

<!-- grade_id: determinant_test, points: 1, grade: true, solution: false -->
```python
arr1 = np.array([[1, 2], [3, 4]])
assert np.isclose(determinant(arr1), -2)

arr2 = np.array([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12]])
assert determinant(arr2) == None

arr3 = np.array([[1]])
assert np.isclose(determinant(arr3), 1)

arr2 = np.array([[1, 0, 0], [0, 1, 0], [0, 0, 1]])
assert np.isclose(determinant(arr2), 1)

arr3 = np.array([[1, 2, 3], [4, 5, 6], [7, 8, 9]])
assert np.isclose(determinant(arr3), 0)

arr5 = np.array([[1, 0], [0, 1], [0, 0]])
assert determinant(arr5) == None

arr6 = np.array([])
assert determinant(arr6) == None

arr8 = np.array([[1, 2], [3, 4], [5, 6]])
assert determinant(arr8) == None

arr9 = np.array([[1, 2, 3], [4, 5, 6]])
assert determinant(arr9) == None

arr10 = np.array([[1, 2], [3, 4], [5, 6]])
assert determinant(arr10) == None
```

# Задание на производные

---
## Задание 8 (1 балл)

Напишите функция для вычисления производной функции $\frac{7x^3}{3} - \frac{5x^2}{2}$ в точке x. Функция должна возврвращать переменную df_dx

С этим заданием возникли сложности, код получается запускать вне функции, производная считается, а вот как обернть это в функцию у меня вопрос. Пробовал по всякому, но не вышло.

```python
dy_dx = 0
def calculate_dderivative(x: float):

    # НАЧАЛО ВАШЕГО РЕШЕНИЯ
    x = sp.Symbol('x')
    
    y = ((7*(x**3))/(3)) - ((5*(x**2))/(2))

    dy_dx = y.diff(x)
    
    # КОНЕЦ ВАШЕГО РЕШЕНИЯ
    return dy_dx
```

<!-- grade_id: calculate_dderivative_test, points: 1, grade: true, solution: false -->
```python
assert np.isclose(calculate_dderivative(0), 0, atol=0.1)
assert np.isclose(calculate_dderivative(1), 2, atol=0.1)
assert np.isclose(calculate_dderivative(2), 18, atol=0.1)
assert np.isclose(calculate_dderivative(3), 48, atol=0.1)
assert np.isclose(calculate_dderivative(5), 150, atol=0.1)
```

Output:
```
File <__array_function__ internals>:200, in isclose(*args, **kwargs)

File c:\Users\Slava\AppData\Local\Programs\Python\Python311\Lib\site-packages\numpy\core\numeric.py:2377, in isclose(a, b, rtol, atol, equal_nan)
   2374     dt = multiarray.result_type(y, 1.)
   2375     y = asanyarray(y, dtype=dt)
-> 2377 xfin = isfinite(x)
   2378 yfin = isfinite(y)
   2379 if all(xfin) and all(yfin):

TypeError: ufunc 'isfinite' not supported for the input types, and the inputs could not be safely coerced to any supported types according to the casting rule ''safe''
```